```python
# Initialize Otter
import otter
grader = otter.Notebook("logistic_regression_assignment_ling_cs_581_2026_submission_eight.ipynb")
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("logistic_regression_assignment_ling_cs_581_2026.ipynb")
```

**Q1:**  The code  given in the code help notebook constructs a toy corpus derived from the Brown corpus and trains a logistic regression language model `lrc` on that corpus.   

Use the trained model `lrc` to predict the probabilities of the entire training set and use those to to **evaluate** the  model by computing its **perplexity** with respect to the training data.  

You should use the definition of perplexity given in our textbook ngrams chapter (see the slides).  Assign the perplexity value to `perp1.`

```python
perp1 = 37.206005418592554
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Q2**:  How many features does the model evaluated in Q1 use?  Assign the answer to `num_features1`.

```python
num_features1 = 508
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Q3**: The code  given in the code help notebook calls a function named `prepare_korpus` to construct the 2D array and predicted word sequence (`Y`) used to train the classifier `lrc`.  Modify the call to that function to build a new `korpus` and use it to train a model which does not use trigger words. (Note:  **You do not have to modify the function, just change how it is called.**) The resulting model only uses the previous word to predict the next word.  It is in effect a classic bigram model. Use the bigram model to predict the probabilities of the entire training set and use those to to **evaluate** this new model by computing its **perplexity** with respect to the training data.  Assign the perplexity value to `perp_bigram`. If the model that uses trigger words is better than the bigram model, assign the value `True` to `trigger_model_is_better`.  Otherwise assign the value `False`.

```python
perp_bigram = 44.168512432090985
trigger_model_is_better = True
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Q4**: How many features does the model evaluated in Q3 use? Assign the answer to num_features_bigram.

```python
num_features_bigram = 358 
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Q5**:  Compute the perplexity the model would assign to the corpus if it had no information, that is, if it predicted every  word, including the correct word, with probability `p`, where `p` is uniformly distributed among the target words.   Assign this perplexity to `perp_uniform`. Two notes.  First: the set of **target words**, words that are candidate predictions,  is not the same as the vocabulary, since only a subset of the vocabulary is found in the training data as a possible target word.  Find the `targets` subset and compute your answer accordingly.  Second:  You do not need to train another model to answer this question (and you shouldn't).

```python
perp_uniform = 91
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Q6**:  Of all the target words in the data, which is predicted with the greatest **precision**? Use the model with trigger words.  You will need to look at the example using `sklearn.metrics.precision_score` in the code help notebook. You are interested in computing per-class precision scores. Then find the target word with the maximal precision score and assign it to `max_precision_word`.

```python
max_precision_word = "time"
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

 This is a gradescope assignment.
 Write your code and compute your values in the code help NB.
 Then assign those values to the appropriate variables in the solution cell of each question in this NB.
 Upload this notebook file ('.ipynb') with your solutions to Canvas.
 Save your code help NB for discussion, but you do not need to hand it in.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 4)
```
>>> 37 < perp1 < 38
True
```


Tests for q2 (points: 2)
```
>>> 500 < num_features1 < 550
True
```


Tests for q3 (points: 4)
```
>>> 44 < perp_bigram < 45
True
```
```
>>> trigger_model_is_better
True
```


Tests for q4 (points: 4)
```
>>> 300 < num_features_bigram < 400
True
```


Tests for q5 (points: 4)
```
>>> 90 < perp_uniform
True
```


Tests for q6 (points: 2)
```
>>> len(max_precision_word) == 4
True
```
