```python
# Initialize Otter
import otter
grader = otter.Notebook("ngram_assignment.ipynb")
```

Output:
```
/usr/local/lib/python3.12/dist-packages/otter/check/utils.py in event_logger(wrapped, self, args, kwargs)
    236         """
    237         try:
--> 238             ret: Optional[LoggedEventReturnValue[T]] = wrapped(*args, **kwargs)
    239 
    240         except Exception as e:

/usr/local/lib/python3.12/dist-packages/otter/check/notebook.py in __init__(self, nb_path, tests_dir, tests_url_prefix)
    105 
    106         if self.interpreter is IPythonInterpreter.COLAB and not os.path.isdir(tests_dir):
--> 107             raise ValueError(f"Tests directory {tests_dir} does not exist")
    108 
    109         cls = type(self)

ValueError: Tests directory ./tests does not exist
```

**Question 1:**  Use the ngram assignment code help notebook to find the necessary unigram and bigram counts to estimate $P(\text{the}\mid \text{of})$ in the  lower-cased version of the Brown corpus. Assign the unigram count to `unigram_ct` and the bigram count to `bigram_ct`. Assign the probability value to `p_the_given_of`.

```python
import nltk
from nltk.corpus import brown
from nltk import download
download('brown')
```

Output:
```
[nltk_data] Downloading package brown to /root/nltk_data...
[nltk_data]   Unzipping corpora/brown.zip.
True
```

```python
bw = brown.words()
bw_lower = [w.lower() for w in bw]
fd = nltk.FreqDist(bw_lower)
bigrams = nltk.bigrams(bw_lower)
bigram_fd = nltk.FreqDist(bigrams)
```

```python
unigram_ct = fd['of']
bigram_ct = bigram_fd['of', 'the']
p_the_given_of = bigram_ct/unigram_ct
# print(p_the_given_of)
```

```python
grader.check("q1")
```

**Question 2:**  Use the ngram assignment code help notebook to find the necessary unigram and bigram counts to estimate $P(\text{player}\mid \text{football})$  in the lower-cased version of the Brown corpus. Assign the unigram count to `unigram_ct_2` and the bigram count to `bigram_ct_2`. Assign the probability value to `p_player_given_football`. Next, find any additional counts you need to estimate the probability of the bigram *football player*. Assign the bigram probability to `p_football_player`.

```python
bigram_ct_2 = bigram_fd['football', 'player']
unigram_ct_2 = fd['football']
p_player_given_football = bigram_ct_2/unigram_ct_2
p_football_player = bigram_fd['football', 'player'] / bigram_fd.N()
print(p_player_given_football)
print(p_football_player)
```

Output:
```
0.027777777777777776
8.611847663304314e-07
```

```python
grader.check("q2")
```

**Question 3:**  This question refers to `p_player_given_football` and `p_football_player` computed in question 2.  According to the **Chain Rule**, dividing one of these probabilities by the other should give a third probability. Call that third probability `p_chained`.  You can estimate `p_chained` independently using counts from the Brown corpus. This will allow you to check that your work in answering question 2 agrees with the Chain Rule.   Assign one of these counts to `count_1` and the other to `count_2`; assign `count_1` divided by `count_2` to `p_chained_2`; `p_chained_2` should  approximately equal (difference less than `1e11`) `p_chained`.

Note:  To facilitate automatic grading, do not use names assigned values in previous questions in answering this one.  Your count assignments should have integers on the right hand side.

```python
p_chained = p_football_player/p_player_given_football
count_1 = fd['football']
count_2 = fd.N()
p_chained_2 = count_1/count_2
```

```python
grader.check("q3")
```

**Question 4:**  Given that a word token follows *of*, how likely is it to be one of the 20 most common words to follow *of*.  In other words, find the sum of the conditional probabilities $P(\text{x}\mid \text{of})$ for the 20 x most likely to follow the word *of*.  Assign the value to `top_20_of_followers`.

```python
import numpy as np

def restrict_most_common (fd,f_rest):
    """
    Returns the x that satisfy restriction f_rest in sorted order
    according to fd.most_common()
    """
    bs,cts = zip(*[(x,ct) for (x,ct) in fd.most_common() if f_rest(x)])
    return bs, np.array(cts)

bs, cts = restrict_most_common(bigram_fd, lambda x: x[0] == 'of')
```

```python
total_of_ct = sum(cts)
sum_of_cts_of_top_20_of_followers = sum(cts[:20])
top_20_of_followers = sum_of_cts_of_top_20_of_followers/total_of_ct
```

```python
grader.check("q4")
```

**Question 5:**  Using the lower-cased version of the Brown corpus, find the most common trigram that contains neither a punctuation mark nor the word "of".  Assign the value to `most_common_lexical_trigram`. Hint: You are allowed to import more help from `nltk` and you may use `string.punctuation` (`string` is a module in the standard Python libraries).

```python
import string
trigrams = list(nltk.trigrams(bw_lower))
trigrams_fd = nltk.FreqDist(trigrams)
no_no_words = list(string.punctuation) + ['of']
restriction = lambda x: all([w not in no_no_words for w in x])
ts, tri_cts = restrict_most_common(trigrams_fd, restriction)
most_common_lexical_trigram  = ts[0]
```

```python
grader.check("q5")
```

**Question 6**:  Suppose in the midst of a deadly worldwide plague your research group believes it has developed a test allowing early detection of those who have contracted the virus.  You test 100 random selected subjects and then quarantine them for 14 days to see which ones have the virus.  Using $\hat{+}$ for those who have tested positive and $+$ for those who actually have the virus, and $\hat{-}$ for those who have tested negative and $-$ for those who are actually free of the virus, here are the results expressed in the form of what is called a **confusion matrix**:

$$
\begin{array}{|c|c|c|}
\hline
         &  +  &  -  \\
         \hline
\hat{+}  &  30 &  40 \\
\hat{-}  &  20 &  10\\
\hline
\end{array}
$$

Answer the following questions

(a) Compute $P(\hat{+}\mid +)$  and assign the value to $p0$. Compute the value of  $P(+ \mid \hat{+})$ and assign the value to $p1$. Show your work.

(b)  Look up the definitions of **precision** and **recall**.  Do either of the probabilities in (a)  measure the precision of the test?  If so, set `precison` to that value.  Otherwise, set it to `None`.  Do either of these probabilities measure the recall of the test?  If so, set `recall` to that value.  Otherwise, set it to `None`.

(c) Compute the maximum likelihood estimate of  $P(+)$ using this sample and set it to $p2$.

(d) Determine which is more likely for this test, a false outcome for a positive test result or a false outcome for a negative test result.  Set `p_f_negative` to the probability of false outcome given a negative test outcome and `p_f_positive` to the probability of a false outcome given a positive test result.  Show your work. Does either of these probabilities show the test is informative in some way?

```python
import numpy as np

p0 = 30/50
p1 = 30/70
p2 = 30 / (30+20)
recall = p0
precision = p1
# p3 = ...
p_f_negative = 20/30
p_f_positive = 40/70
```

```python
grader.check("q6")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

 This is a gradescope assignment.
 Write your code and compute your values in the code help NB.
 Then assign those values to the appropriate variables in the solution cell of each question in this NB.
 Upload this notebook file ('.ipynb') with your solutions to Canvas.
 Save your code help NB for discussion, but you do not need to hand it in.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 2)
```
>>> def test_of_ct(unigram_ct):
...     return unigram_ct == 36412
>>> test_of_ct(unigram_ct)
True
```
```
>>> def test_of_the_ct(bigram_ct):
...     return bigram_ct == 9717
>>> test_of_the_ct(bigram_ct)
True
```
```
>>> def test_p_the_given_of(p_the_given_of):
...     return 0.26686 < p_the_given_of < 0.26687
>>> test_p_the_given_of(p_the_given_of)
True
```


Tests for q2 (points: 4)
```
>>> def test_bigram_ct_2(bigram_ct_2):
...     return bigram_ct_2 == 1
>>> test_bigram_ct_2(bigram_ct_2)
True
```
```
>>> def test_unigram_ct_2(unigram_ct_2):
...     return unigram_ct_2 == 36
>>> test_unigram_ct_2(unigram_ct_2)
True
```
```
>>> def test_p_football_player(p_football_player):
...     return 0 < p_football_player < 8.7e-07
>>> test_p_football_player(p_football_player)
True
```
```
>>> def test_p_player_given_football(p_player_given_football):
...     return 0.0275 < p_player_given_football < 0.028
>>> test_p_player_given_football(p_player_given_football)
True
```


Tests for q3 (points: 2)
```
>>> count_1 == 36
True
```
```
>>> count_2 == 1161192 or count_2 == 1161191
True
```
```
>>> 3.10026248889e-05 < p_chained_2 < 3.1002624889e-05
True
```


Tests for q4 (points: 4)
```
>>> def test_of_followers(top_20_of_followers):
...     return 0.4427 < top_20_of_followers < 0.4428
>>> test_of_followers(top_20_of_followers)
True
```


Tests for q5 (points: 4)
```
>>> most_common_lexical_trigram in [('the', 'united', 'states'), 'the united states', ('The', 'United', 'States'), 'The United States', ('the', 'United', 'States'), 'the United States']
True
```


Tests for q6 (points: 4)
```
>>> p0 == 0.6
True
```
```
>>> 0.4285 < p1 < 0.4286
True
```
```
>>> precision == p1
True
```
```
>>> recall == p0
True
```
```
>>> 0.6666 < p_f_negative < 0.6667
True
```
```
>>> 0.5714 < p_f_positive < 0.5715
True
```
